# Tree-based models - Exercises

```python
import warnings
warnings.filterwarnings("ignore")

import hashlib

import pandas as pd
import numpy as np

from utils.utils import *
```

# Exercise 1 - Decision Trees

## 1.1 Gini impurity

Used by the CART algorithm for classification, Gini impurity is an alternative to entropy.

As entropy, it is a way to measure node homogeneity. Hence, it can be used to identify promissing splits.

Take $p$ as the probability of the positive class, i.e., the proportion of positive cases in the set:

$$I_G(p)= 1 - p^2 - (1-p)^2$$

It measures how often a randomly chosen element from the set would be incorrectly labeled.

<!-- grade_id: cell-096c9f233958084b, points: 0, grade: false, solution: true -->
```python
def gini(p):
    # YOUR CODE HERE
    return 1-p**2-(1-p)**2
    
    
data = make_data()
data.head(20)
```

Output:
```
Outlook Temperature Humidity  Windy  Class
0      sunny         hot     high  false      0
1      sunny         hot     high   true      0
2   overcast         hot     high  false      1
3       rain        mild     high  false      1
4       rain        cool   normal  false      1
5       rain        cool   normal  false      0
6   overcast        cool   normal  false      1
7      sunny        mild     high  false      0
8      sunny        cool   normal  false      1
9       rain        mild   normal  false      1
10     sunny        mild   normal   true      1
11  overcast        mild     high   true      1
12  overcast         hot   normal  false      1
13      rain        mild     high   true      0
```

<!-- grade_id: cell-cd65813e62272b59, points: 4, grade: true, solution: false -->
```python
expected_hash_1 = '5feceb66ffc86f38d952786c6d696c79c2dbc239dd4e91b46729d73a27fb57e9'
assert hashlib.sha256(str(gini(p=0)).encode('utf-8')).hexdigest() == expected_hash_1

expected_hash_2 = '0061765517977143a7ed83b494b172f02af75de1c3cdde6f969effdef053ee8c'
assert hashlib.sha256(str(gini(p=1/6)).encode('utf-8')).hexdigest() == expected_hash_2

expected_hash_3 = '683487f42b8f523b6df156fbc42818868de8d2b70ad48f801df4df36d07b40f2'
assert hashlib.sha256(str(gini(p=1/3)).encode('utf-8')).hexdigest() == expected_hash_3

expected_hash_4 = 'd2cbad71ff333de67d07ec676e352ab7f38248eb69c942950157220607c55e84'
assert hashlib.sha256(str(gini(p=1/2)).encode('utf-8')).hexdigest() == expected_hash_4
```

## 1.2 Applying the Gini 

### 1.2.1 Single node

Compute the impurity of a node that includes all instances where $x_i^{Windy}$ is `'true'`.

(Note that `'true'` is a string.)

```python

```

Output:
```
Outlook Temperature Humidity Windy  Class
1      sunny         hot     high  true      0
10     sunny        mild   normal  true      1
11  overcast        mild     high  true      1
13      rain        mild     high  true      0
```

<!-- grade_id: cell-1eb99a6c8f5fd94d, points: 0, grade: false, solution: true -->
```python
single_node_gini = data[data.Windy=='true']

# single_node_gini = ...
def compute_probability(node):
    # YOUR CODE HERE
    return node['Class'].sum()/node['Class'].count()

single_node_gini = gini(compute_probability(single_node_gini))
# YOUR CODE HERE
single_node_gini
```

Output:
```
0.5
```

<!-- grade_id: cell-f847f2e82a8bd209, points: 4, grade: true, solution: false -->
```python
expected_hash_5 = 'd2cbad71ff333de67d07ec676e352ab7f38248eb69c942950157220607c55e84'
assert hashlib.sha256(str(single_node_gini).encode('utf-8')).hexdigest() == expected_hash_5
```

```python
data['Outlook'].unique()
```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-15-404194bcbadc> in <module>
----> 1 data['Outlook'].unique().values

AttributeError: 'numpy.ndarray' object has no attribute 'values'
```

### 1.2.2 Single feature

Write a function to compute the mean impurity of branching on a given feature.

(Hint: compute the impurity of each resulting node and average the results.)

<!-- grade_id: cell-ca4c34be595ebcb3, points: 0, grade: false, solution: true -->
```python
from statistics import mean


def mean_impurity(data, feature_name):
    # YOUR CODE HERE
    impurities = []
    for split in data[feature_name].unique():

        single_node = data[data[feature_name]==split]

        impurities.append( gini(compute_probability(single_node)))
    return np.mean(impurities)
```

<!-- grade_id: cell-92d7465bbb671305, points: 4, grade: true, solution: false -->
```python
assert mean_impurity(data, 'Temperature') == 0.4398148148148148

expected_hash_6 = '60bb8f6c52a2aca85b1a3ae08e71f97dd3fa8603b55abefb09ebb0edfe5294d1'
assert hashlib.sha256(str(mean_impurity(data, 'Outlook')).encode('utf-8')).hexdigest() == expected_hash_6

expected_hash_7 = '1bd499c4b5bbae82a4f4a4efe3950b803e7dc64764c56115810e696db1b28afe'
assert hashlib.sha256(str(mean_impurity(data, 'Humidity')).encode('utf-8')).hexdigest() == expected_hash_7

expected_hash_8 = 'ccb9ffbc97edf0207e9b20dd9331f576984b45f178aa756421dd16508dd984fc'
assert hashlib.sha256(str(mean_impurity(data, 'Windy')).encode('utf-8')).hexdigest() == expected_hash_8
```

## 1.3 Worst split

Use the `mean_impurity()` function above to identify the **least promising** feature.

(Uncomment the **correct option**.)

```python
for c in data.columns:
    print('{}: {}'.format(c,mean_impurity(data, c)))
```

Output:
```
Outlook: 0.32
Temperature: 0.4398148148148148
Humidity: 0.3673469387755103
Windy: 0.45999999999999996
Class: 0.0
```

<!-- grade_id: cell-1adca6f8556b71f2, points: 0, grade: false, solution: true -->
```python
# worst_split = 'Temperature'
#worst_split = 'Outlook'
# worst_split = 'Humidity'
worst_split = 'Windy'

# YOUR CODE HERE

```

<!-- grade_id: cell-2b25486151a99ff4, points: 4, grade: true, solution: false -->
```python
expected_hash_9 = 'cdcaccab2e98e50c7fd550e5be166e6c260312708903794361d1815222ed0975'
assert hashlib.sha256(str(worst_split).encode('utf-8')).hexdigest() == expected_hash_9
```

# Exercise 2 - Random Forests

## 2.1 Bagging

Uncomment the **incorrect statement**.

<!-- grade_id: cell-c734638ada0323ab, points: 0, grade: false, solution: true -->
```python
# incorrect_statement_1 = 'Bagging helps to deal with overfitting, which is a big risk when using Decision Trees.'
# incorrect_statement_1 = 'Bagging involves averaging the predictions of multiple independent models.'
# incorrect_statement_1 = 'Bagging involves creating multiple data sets by sampling observations.'
incorrect_statement_1 = 'Bagging involves creating multiple data sets by sampling columns.'

# YOUR CODE HERE

```

<!-- grade_id: cell-f0ee1005cf77cd42, points: 1, grade: true, solution: false -->
```python
expected_hash_10 = '1fb9060eefe647e998b5868f95d4930673fa3645ed0120176f6867f6c43ecf71'
assert hashlib.sha256(str(incorrect_statement_1).encode('utf-8')).hexdigest() == expected_hash_10
```

## 2.2 Random forests

Uncomment the **incorrect statement**.

<!-- grade_id: cell-7841ec3e8a3d8ab4, points: 0, grade: false, solution: true -->
```python
# incorrect_statement_2 = 'Random forests rely on bootstrapping to build several training sets.'
# incorrect_statement_2 = 'Random forests rely on random feature selection before each split.'
# incorrect_statement_2 = 'Random forests aggregates the predictions of multiple models running in parallel.'
incorrect_statement_2 = 'Random forests aggregates the predictions of multiple models running sequentially.'

# YOUR CODE HERE

```

<!-- grade_id: cell-d7e4db2d6899e82f, points: 1, grade: true, solution: false -->
```python
expected_hash_11 = '3a36af9faaf10138801ffef535b24bfe0a34217160a886dacdb3bd68b7777067'
assert hashlib.sha256(str(incorrect_statement_2).encode('utf-8')).hexdigest() == expected_hash_11
```

# 3 Gradient boosting

Uncomment the **incorrect statement**.

<!-- grade_id: cell-19d101619190adb0, points: 0, grade: false, solution: true -->
```python
# incorrect_statement_3 = 'Gradient boosting fits individual trees sequentially.'
incorrect_statement_3 = 'Gradient boosting fits individual trees on the predictions from the previous tree.'
# incorrect_statement_3 = 'Gradient boosting fits individual trees on the pseudo-residuals of the previous tree.'
# incorrect_statement_3 = 'Gradient boosting can be used to optimize a differentiable cost function.'

# YOUR CODE HERE

```

<!-- grade_id: cell-2aa78ad9cfadcba8, points: 2, grade: true, solution: false -->
```python
expected_hash_12 = '4000dc46debb11c42c4d44c82eb50120463cad49ae182b5130c9e276104570bd'
assert hashlib.sha256(str(incorrect_statement_3).encode('utf-8')).hexdigest() == expected_hash_12
```

```python

```