Yay! Exercises! 

```python
import pandas as pd 
from matplotlib import pyplot as plt 
from sklearn.linear_model import LinearRegression
import numpy as np
import hashlib # for grading purposes
%matplotlib inline 
import utils
```

### You're hired

You've been hired as the data scientist for a supermarket chain that wants to start exctracting insights from their data. First you'll start with a specific store.

Let's get our data:

```python
store = utils.get_store_data()
store.head()
```

Output:
```
date  customers
67528  2016-10-18       1655
55251  2016-02-27       1314
49000  2015-10-28       1784
54244  2016-02-08        417
48311  2015-10-15       1725
```

#### Q1: make the index the datetime of the dates. 

<!-- grade_id: cell-c5dde6c9d4e1e881, points: 0, grade: false, solution: true -->
```python
# We expect the solution to be a dataframe

store['date'] = pd.to_datetime(store.date, infer_datetime_format=True)
store = store.set_index('date',    # <---- Set the index to be our timestamp data  
                      drop=True)
store = store.sort_index()           # Don't forget best practices

# YOUR CODE HERE
```

<!-- grade_id: cell-018ff5132531d38c, points: 1, grade: true, solution: false -->
```python
expected_hash = '24e6fe346ad970255e2da6db24d4721d0a240cbb5a00a58fe731b4ad767d6ab0'
assert hashlib.sha256(str(store.iloc[25].name).encode()).hexdigest() == expected_hash
assert hashlib.sha256(str(store.index.dtype).encode()).hexdigest() == '261738f2e43a1c47a16f043b46deb993943d61f4a2bbe5ef4b03c3fb1af362b5'


# clue: if this assert is failing, and your iloc[25].name is '18/02/2015', 
# then you are missing the "best practices" part. 
# What did we say in the Learning notebook about this? 
```

#### Q2: what was the worse day in terms of customers on December 2016?

```python

```

<!-- grade_id: cell-f3ba0f907ca7e1b7, points: 0, grade: false, solution: true -->
```python
# hint: the answer should be a timestamp

worst_day = store.customers['December 2016'].idxmin()

# YOUR CODE HERE
```

<!-- grade_id: cell-141eb4112e911043, points: 2, grade: true, solution: false -->
```python
expected_hash = 'a857d04e7b105cd2c922665d08abcba6474f8adbc46d68f34ad851ed99ae6b40'
assert hashlib.sha256(str(worst_day).encode()).hexdigest() == expected_hash
```

#### Q3: last wednesday there were 2800 customers, your boss said he's never seen so much people on a wednesday. To confirm if it is true, can you find the maximum number of customers that we've ever had on a wednesday?

- _hint #1: you can use the methods at the bottom of this [page](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DatetimeIndex.html) directly on the index_  
- _hint #2: when operating directly on the index, you do not need to use `.dt` to use the methods_

```python
store.index.day_name()
```

Output:
```
Index(['Wednesday', 'Thursday', 'Friday', 'Saturday', 'Sunday', 'Monday',
       'Tuesday', 'Wednesday', 'Thursday', 'Friday',
       ...
       'Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday',
       'Saturday', 'Sunday', 'Monday', 'Tuesday'],
      dtype='object', name='date', length=1676)
```

<!-- grade_id: cell-089cdded3864f17b, points: 0, grade: false, solution: true -->
```python
# max_customers_wednesday =
max_customers_wednesday = store[store.index.day_name() == 'Wednesday'].max().customers
# YOUR CODE HERE
```

<!-- grade_id: cell-6d882e41cf9413ae, points: 2, grade: true, solution: false -->
```python
expected_hash = '99dbbc4aad5da980a3b5035ff28d2e9f828d973cf7be91cdf3b6886741879371'
assert hashlib.sha256(str(max_customers_wednesday).encode()).hexdigest() == expected_hash
```

#### Q4: "today there were news about an approaching zombie apocalypse, everyone came to buy stuff. I've never seen such an absolute increase in sales from one day to the next. Yesterday there were 100 customers, today there were 5000." To confirm if what your boss is saying is true, can you find the maximum increase in sales from one day to the next?

<!-- grade_id: cell-f19f56a68b7febcb, points: 0, grade: false, solution: true -->
```python
# hint: the solution expects a float

max_increase = store.diff().max().customers

# YOUR CODE HERE
```

<!-- grade_id: cell-0703bbeda85d0834, points: 2, grade: true, solution: false -->
```python
expected_hash = '879f277341e3578a03df02d7f9a1e98da6b7678e09b8490db4fa87a64b77c414'
assert hashlib.sha256(str(max_increase).encode()).hexdigest() == expected_hash
```

#### Q5: It turned out the zombie apocalypse was fake news, but it still was a week with a signficant total amount of customers: 13000. What was the last week we had more transactions than this?

<!-- grade_id: cell-3745ea4d56722003, points: 0, grade: false, solution: true -->
```python
# We expect the answer to be a weekly time period (freq='W-SUN'), so answers  
# which are of the form "the first monday of that week" will not pass the grader. 
# hint: by default pandas uses freq='W-SUN'

#total_weekly_customers = 
#last_week_customers = 

# YOUR CODE HERE
total_weekly_customers = store.resample('W-SUN').sum()
total_weekly_customers = total_weekly_customers[total_weekly_customers.customers>=13000]
last_week_customers = total_weekly_customers.index.max()
```

<!-- grade_id: cell-e8248641ed113328, points: 3, grade: true, solution: false -->
```python
expected_hash = '392462e1c962b89eb0b84a245f954a54f7d3c7ac930508ef8df909ee02586613'
assert hashlib.sha256(str(last_week_customers).encode()).hexdigest() == expected_hash
```

### Congrats!

Your work is proving useful, so your boss has asked to expand your analysis to the whole chain. 

Let's get the new data:

```python
stores = utils.get_stores_data()
stores.head()
```

Output:
```
date  store_nbr  customers
43273  2015-07-09          7       1830
18371  2014-02-02          6       2210
32615  2014-12-01         10       1107
62678  2016-07-16         10       1002
58823  2016-05-04         24       2306
```

```python
data = utils.get_stores_data()
print('We now have %0.0f data points. Wooooow!' % len(data))
```

Output:
```
We now have 83488 data points. Wooooow!
```

The thing is, we can't just set the index to be the day, as we now have multiple stores on the same day. 

Looks like we have to go into multi-indexing...

#### Q6: make the date into a datetime, and then set the index to be `[Date, Store]`

<!-- grade_id: cell-76e9757d88ddc4b1, points: 0, grade: false, solution: true -->
```python
stores['date'] = pd.to_datetime(stores['date'])
stores = stores.set_index(['date','store_nbr'])
stores = stores.sort_index()

# YOUR CODE HERE

```

<!-- grade_id: cell-90a9c56f23c1116b, points: 1, grade: true, solution: false -->
```python
assert hashlib.sha256(str(stores.index.get_level_values(0)).encode()).hexdigest() ==  \
        'aff1142eb44c03cc16656f3c3541ade7b5c151abcb4340fb14a058a7fec7e46b'
assert hashlib.sha256(str(stores.index.get_level_values(1)).encode()).hexdigest() ==  \
        'd28cc6e9308dd2e399e100566cc5366b0ef79b74d84a5d666d792c4514704bfb'
assert hashlib.sha256(str(stores.index.get_level_values(0).dtype).encode()).hexdigest() == \
        '261738f2e43a1c47a16f043b46deb993943d61f4a2bbe5ef4b03c3fb1af362b5'
```

#### Q7: when did store 20 open?

```python
stores.head()
```

Output:
```
customers
date       store_nbr           
2013-01-01 25               770
2013-01-02 1               2111
           2               2358
           3               3487
           4               1922
```

```python
idx = pd.IndexSlice
```

```python

```

<!-- grade_id: cell-bc9fe27354d7f35b, points: 0, grade: false, solution: true -->
```python
# hint: the answer should be a timestamp


store_20_open=stores.loc[idx[:,20],:].index.get_level_values(0).min()

# YOUR CODE HERE

```

<!-- grade_id: cell-ec809d4ad7cf82c7, points: 2, grade: true, solution: false -->
```python
expected_hash = '789fe6da1fa44792a15fa1285ef9a11021bf3f7771255545462153f69338b3c9'
assert hashlib.sha256(str(store_20_open).encode()).hexdigest() == expected_hash
```

#### Q8: how many stores opened in 2015?

```python
stores.loc[idx['2015',:],:].index.get_level_values(1).nunique()
```

Output:
```
53
```

<!-- grade_id: cell-ff6e9c588b6ab2bf, points: 0, grade: false, solution: true -->
```python
# hint: you can assume that no stores closed in 2015.

nr_stores_2014 = len(stores.loc[idx['2014',:],:].groupby('store_nbr'))
nr_stores_2015 = len(stores.loc[idx['2015',:],:].groupby('store_nbr'))
nr_stores_opened_2015 = nr_stores_2015-nr_stores_2014
nr_stores_opened_2015
# YOUR CODE HERE
```

Output:
```
5
```

<!-- grade_id: cell-85511006b783f6f0, points: 2, grade: true, solution: false -->
```python
expected_hash = 'ef2d127de37b942baad06145e54b0c619a1f22327b2ebbcfbec78f5564afe39d'
assert hashlib.sha256(str(nr_stores_opened_2015).encode()).hexdigest() == expected_hash
```

#### Q9: Does the total number of customers that go to stores have a stronger weekly or bi-weekly seasonality during 2016?

```python
stores.groupby('date').sum().head(5)
```

Output:
```
customers
date                 
2013-01-01        770
2013-01-02      93215
2013-01-03      78504
2013-01-04      78494
2013-01-05      93573
```

<!-- grade_id: cell-90ce7e7916ff582b, points: 0, grade: false, solution: true -->
```python
# We expect you to use correlation as a proxy for seasonality.

# hint 1: Get the total sales per day (sum over all stores), and see whether seasonality 
# is bigger for 7 or for 14 days.

# hint2: the order suggested below has a reason. You can think about what would 
# happen if you did it in the opposite order.

# hint3: we expect you to use correlation as a proxy for seasonality. 
# the grader expects you to use the pandas method for correlation

# hint4: use a negative shift

total_sales_per_day = stores.groupby('date').sum()

total_sales_per_day_week_before = total_sales_per_day.shift(-7)
total_sales_per_day_2_weeks_before = total_sales_per_day.shift(-14)

total_sales_per_day_2016_week_before = total_sales_per_day_week_before['2016']
total_sales_per_day_2016_2_weeks_before = total_sales_per_day_2_weeks_before['2016']

weekly_corr = total_sales_per_day.customers.corr(total_sales_per_day_2016_week_before.customers)
#biweekly_corr = 
weekly_corr
# YOUR CODE HERE
```

Output:
```
0.4260724637683414
```

<!-- grade_id: cell-f44cea818496f41b, points: 2, grade: true, solution: false -->
```python
expected_hash1 = '96cc11d2ee80bddd4e7241aa7b6abac54810adc28140198c1525a4f1ba9a78db'
expected_hash2 = '2c66c7d6bee05863b57168ff4f235a7811905085e937d492bf379d379cd3ed95'
assert hashlib.sha256(str(weekly_corr).encode()).hexdigest() == expected_hash1
assert hashlib.sha256(str(biweekly_corr).encode()).hexdigest() == expected_hash2
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-27-1293823d0348> in <module>
      1 expected_hash1 = '96cc11d2ee80bddd4e7241aa7b6abac54810adc28140198c1525a4f1ba9a78db'
      2 expected_hash2 = '2c66c7d6bee05863b57168ff4f235a7811905085e937d492bf379d379cd3ed95'
----> 3 assert hashlib.sha256(str(weekly_corr).encode()).hexdigest() == expected_hash1
      4 assert hashlib.sha256(str(biweekly_corr).encode()).hexdigest() == expected_hash2

AssertionError:
```

#### **Q10: What is the store with the strongest positive trend? Only including stores that were open since 2013**

```python
for store in stores.index.get_level_values(1):
    print(store)
```

Output:
```
25
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
23
24
25
26
27
28
30
31
32
33
34
35
37
38
39
40
41
43
44
45
46
47
48
49
50
51
54
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
23
24
25
26
27
28
30
31
32
33
34
35
37
38
39
40
41
43
44
45
46
47
48
49
50
51
54
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
23
24
25
26
27
28
30
31
32
33
34
35
37
38
39
40
41
43
44
45
46
47
48
49
50
51
54
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
23
24
25
26
27
28
30
31
32
33
34
35
37
38
39
40
41
43
44
45
46
47
48
49
50
51
54
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
23
24
25
26
27
28
30
31
32
33
34
35
37
38
39
40
41
43
44
45
46
47
48
49
50
51
54
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
23
24
25
26
27
28
30
31
32
33
34
35
37
38
39
40
41
43
44
45
46
47
48
49
50
51
54
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
23
24
25
26
27
28
30
31
32
33
34
35
37
38
39
40
41
43
44
45
46
47
48
49
50
51
54
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
23
24
25
26
27
28
30
31
32
33
34
35
37
38
39
4
```

<!-- grade_id: cell-3eba3fa23da4ac51, points: 0, grade: false, solution: true -->
```python
# This one is a bit harder and there are several ways to solve it. For this reasons the
# body of the solution is not presented as in the previous questions. 

# A few clues:
# don't forget to select only the stores that were open since 2013. 
# most solutions are going to require a for loop to run through the various stores
# assume linear trends, where the trend is characterized by the slope of the linear regression



# max_trend_store = 

# YOUR CODE HERE
coefs = []
stores_s2013 = stores[stores.index.get_level_values(0)>='2013']
for store in stores.index.unique(level=1):
    curr_store = stores_s2013.loc[idx[:, store], :]
    curr_store = curr_store.reset_index()
    X = curr_store.reset_index().index.values.reshape(-1, 1)  
    print(curr_store)
    slr = LinearRegression(fit_intercept=True)
    
    slr.fit(X, curr_store['customers'])
    coefs.append([store , slr.coef_[0]])

max_trend_store = pd.DataFrame(coefs).sort_values(1).iloc[-1,0]  

```

Output:
```
date  store_nbr  customers
0    2013-01-01         25        770
1    2013-01-02         25       1038
2    2013-01-03         25        887
3    2013-01-04         25       1054
4    2013-01-05         25       1355
5    2013-01-06         25        716
6    2013-01-07         25        703
7    2013-01-08         25        626
8    2013-01-09         25        755
9    2013-01-10         25        632
10   2013-01-11         25        979
11   2013-01-12         25       1368
12   2013-01-13         25        748
13   2013-01-14         25        664
14   2013-01-15         25        644
15   2013-01-16         25        775
16   2013-01-17         25        640
17   2013-01-18         25        971
18   2013-01-19         25       1494
19   2013-01-20         25        816
20   2013-01-21         25        754
21   2013-01-22         25        692
22   2013-01-23         25        795
23   2013-01-24         25        691
24   2013-01-25         25       1041
25   2013-01-26        
```

```python

```

<!-- grade_id: cell-0b050707a83767ed, points: 3, grade: true, solution: false -->
```python
expected_hash = '0e17daca5f3e175f448bacace3bc0da47d0655a74c8dd0dc497a3afbdad95f1f'
assert hashlib.sha256(str(max_trend_store).encode()).hexdigest() == expected_hash
```

```python

```

```python

```